Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Rea Rustagi"
COLLABORATORS = ""
```

---

```python
import matplotlib.pyplot as plt

%matplotlib inline
%precision 16
import numpy

```

# Homework 0:  A Test

This notebook contains two short (and hopefully simple) problems to test the homework collection and grading system for APMA4300

## Question 1  

In the cell below, write out in LaTeX the quadratic formula that solves

$$
    x^2 + \alpha x + \beta = 0
$$

<!-- grade_id: cell1448053803160613, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE
$$
x = \frac{-\alpha \pm \sqrt{\alpha^2 - 4 \beta}}{2}
$$

## Question 2

Write a python function that returns the two roots of your quadratic equation given coefficients $\alpha$, $\beta$

This is actually a little trickier than it looks:
*  you can use numpy.sqrt(x) but be careful, it will fail if $x<0$ unless you call it as numpy.sqrt(x,dtype=complex)
*  The tests assume the two roots are sorted (hint: you can use numpy.sort())

<!-- grade_id: cell3621925556728684, points: 0, grade: false, solution: true -->
```python
def myroots(a, b):
    """Compute the two roots of the quadratic equation given by x^2 + ax + b = 0
       and return them as a python list
    """
    # YOUR CODE HERE
    
    roots = []
    
    x1 = (-a + (numpy.sqrt(a**2-(4*b), dtype=complex)))/2
    x2 = (-a - (numpy.sqrt(a**2-(4*b), dtype=complex)))/2
        
    roots.extend([x1, x2])
    roots.sort()
    
    return roots  

    #raise NotImplementedError()
```

<!-- grade_id: cell5655123106444739, points: 5, grade: true, solution: false -->
```python
numpy.testing.assert_allclose(myroots(1.,0.), [-1.,0.])
numpy.testing.assert_allclose(myroots(1.,-2.), [-2.,1.])
numpy.testing.assert_allclose(myroots(-5.0,6.0), [2.,3.])
print('success: passed obvious tests')
```

Output:
```
success: passed obvious tests
```

```python

```